# Selection

## Import pandas

```python
import pandas as pd
```

## Import data

```python
# URL of data
URL = "https://raw.githubusercontent.com/kirenz/datasets/master/height_clean_cols.csv"
```

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python
df = pd.read_csv(URL)

df["gender"] = df["gender"].astype("category")
df['id'] = df['id'].astype(str)
```

## Selection

### Getting [[]]

Selecting a single column with `[[]]`:

- Select the column `height` and save it as a new Pandas dataframe `df_height`

Hint:

```python
df_height = df[[___]]
````

<!-- grade_id: select1, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_height = df[["height"]]
### END SOLUTION
```

<!-- grade_id: c_select1, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert  df_height.columns.to_list() == ['height']
```

Selecting via [], which slices the rows (endpoint is not included) and includes all columns:

```python
df[0:1]
```

Output:
```
name id  height  average_height_parents  gender  number  height_m  \
0  Stefanie  1     162                   161.5  female      42      1.62   

   weight    bmi        date  
0   84.58  32.23  2022-10-08
```

- Select rows 1, 2 and 3 and save it as `df_sr`

<!-- grade_id: select2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_sr = df[1:4]
### END SOLUTION
```

<!-- grade_id: c_select2, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df_sr.iloc[0,0] == 'Peter'
assert df_sr.iloc[2,0] == 'Manuela'
assert len(df_sr) == 3
```

### By label .loc

The `.loc` (location) attribute is the primary access method. The following are valid inputs:

For getting a cross section using a label:

```python
# select row 0
df.loc[[0]]
```

Output:
```
name id  height  average_height_parents  gender  number  height_m  \
0  Stefanie  1     162                   161.5  female      42      1.62   

   weight    bmi        date  
0   84.58  32.23  2022-10-08
```

Selecting on a multi-axis by label:

```python
# only select location at row 0 for column "name"
df.loc[0 , 'name']
```

Output:
```
'Stefanie'
```

```python
# select row 2 to 4 for column "name" (when using .loc endpoints are included)
df.loc[2:4 , 'name']
```

Output:
```
2    Stefanie
3     Manuela
4       Simon
Name: name, dtype: object
```

```python
# select row 2 to 4 for columns "name" and "height" (when using .loc endpoints are included)
df.loc[2:4 , ['name', 'height']]
```

Output:
```
name  height
2  Stefanie     163
3   Manuela     164
4     Simon     164
```

```python
# select all rows for 
df.loc[ : , ["name", "height"]]
```

Output:
```
name  height
0   Stefanie     162
1      Peter     163
2   Stefanie     163
3    Manuela     164
4      Simon     164
5     Sophia     164
6      Ellen     164
7     Emilia     165
8       Lina     165
9      Marie     165
10      Lena     165
11      Mila     165
12       Fin     165
13      Eric     166
14       Pia     166
15      Marc     166
16     Ralph     166
17       Tom     167
18    Steven     167
19   Emanuel     168
```

- Select row 0 and 1 for columns 'name' and 'height' and save it as `df_loc1`

<!-- grade_id: loc, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_loc1 = df.loc[0:1, ["name", "height"]]
### END SOLUTION
```

<!-- grade_id: c_loc, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert  df_loc1.columns.to_list() == ['name', 'height']
assert  len(df_loc1) == 2
assert  df_loc1.iloc[0,0] == "Stefanie"

```

Reduction in the dimensions of the returned object:

```python
df.loc[0, ["name", "height"]]
```

Output:
```
name      Stefanie
height         162
Name: 0, dtype: object
```

For getting a scalar value:

```python
df.loc[[0], "height"]
```

Output:
```
0    162
Name: height, dtype: int64
```

### By position .iloc

Pandas provides a suite of methods in order to get purely integer based indexing. 

Here, the `.iloc` attribute is the primary access method. 

*When using `.iloc`, endpoints are not included.*

```python
df.iloc[0]
```

Output:
```
name                        Stefanie
id                                 1
height                           162
average_height_parents         161.5
gender                        female
number                            42
height_m                        1.62
weight                         84.58
bmi                            32.23
date                      2022-10-08
Name: 0, dtype: object
```

By integer slices:

```python
df.iloc[0:2, 0:2]
```

Output:
```
name id
0  Stefanie  1
1     Peter  2
```

By lists of integer position locations:

```python
df.iloc[[0, 2], [0, 2]]
```

Output:
```
name  height
0  Stefanie     162
2  Stefanie     163
```

For slicing rows explicitly:

```python
df.iloc[1:3, :]
```

Output:
```
name id  height  average_height_parents  gender  number  height_m  \
1     Peter  2     163                   163.5    male      42      1.63   
2  Stefanie  3     163                   163.2  female      42      1.63   

   weight    bmi        date  
1   70.57  26.56  2022-10-08  
2   75.48  28.41  2022-10-08
```

For slicing columns explicitly:

```python
df.iloc[:, 1:3]
```

Output:
```
id  height
0    1     162
1    2     163
2    3     163
3    4     164
4    5     164
5    6     164
6    7     164
7    8     165
8    9     165
9   10     165
10  11     165
11  12     165
12  13     165
13  14     166
14  15     166
15  16     166
16  17     166
17  18     167
18  19     167
19  20     168
```

For getting a value explicitly:

```python
df.iloc[0, 0]
```

Output:
```
'Stefanie'
```

Use `.iloc` to obtain rows 4 to 6 and columns 3 and 4. Save the result as `df_iloc1` 

<!-- grade_id: iloc, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_iloc1 = df.iloc[4:7, 3:5]
### END SOLUTION
```

<!-- grade_id: c_iloc, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert  df_iloc1.columns.to_list() == ['average_height_parents', 'gender']
assert  len(df_iloc1) == 3
assert  df_iloc1.iloc[0,1] == "male"
```