**Wichtige Hinweise**

- Aktivieren Sie die entsprechende Conda-Umgebung bevor Sie beginnen.

- Geben Sie als `NAME` ihr HdM-Kürzel an.

- Ändern Sie **nicht** den Namen der Datei und löschen Sie keine Zellen.

- Bearbeiten Sie alle Zellen mit dem Hinweis <font color='green'> \# YOUR CODE HERE </font> 

- Die Funktion **NotImplementedError()** soll die Abgabe von leeren Zellen verhindern. Löschen Sie die Funktion, sobald Sie in einer dieser Zellen arbeiten. 

- Stellen Sie sicher, dass alles wie erwartet läuft, bevor Sie die Prüfung abgeben: Starten Sie den Kernel neu und führen Sie alle Zellen aus: wählen Sie "Restart" und dann "Run All" 

```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Change column names

## Import pandas

- To load the pandas package and start working with it, import the package. 


- The community agreed alias for pandas is `pd`, so loading pandas as pd is assumed standard practice for all of the pandas documentation:

```python
import pandas as pd
```

## Import data

```python
# URL of data
URL = "https://raw.githubusercontent.com/kirenz/datasets/master/height_unclean.csv"
df = pd.read_csv(URL, sep=";", decimal=',')
df
```

## Change column names

Usually, we prefer to work with columns that have the following proporties:


- no leading or trailing whitespace (`"name"` instead of `" name "`, `" name"` or `"name "`)


- all lowercase (`"name"` instead of `"Name"`)


- now white spaces (`"my_name"` instead of `"my name"`)

### Simple rename

- First, we rename columns by simply using a mapping
- We rename `"Name"` to `"name"` (we want to display errors and don't save the changes for now).


Hint:

```python
df = df.rename(columns={"OLD_NAME": "NEW_NAME"}, errors="raise")
```

<!-- grade_id: rename_name, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_rename_name, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df.loc[0, 'name'] == "Stefanie"
```

```python
df.head(2)
```

- Let`s rename Gender to gender


- Here, we just want to display the result (without saving it).


- Remove the # and run the following code:

```python
# df.rename(columns={"Gender": "gender"}, errors="raise")
```

- This raises an error. 


- Can you spot the problem? Take a look at the end of the error statement and describe the type of error. How could you fix the problem?

<!-- grade_id: c_rename, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Trailing and leading spaces (with regex)

- We use regular expressions to deal with whitespaces


- To change multiple column names in `df` at once, we use the method `df.columns = df.columns.str___` 

- To replace the spaces, we use `.replace()` with `regex=True`

- In the following function, we search for leading (line start and spaces) and trailing (line end and spaces) spaces and replace them with an empty string:

Hint:

replace r"*this pattern*" with empty string r""


```python
df.columns = df.columns.str.replace(r"___ | ___", r"", regex=True)
```

Explanation for *regex* (see also [Stackoverflow](https://stackoverflow.com/a/67466222)):

- we start with `r` (for raw) which tells Python to treat all following input as raw text (without interpreting it)
- "`^`": is line start
- " ": is a white space
- "`+`": some following characters
- "`|`": is or
- "`$`": is line end
- "": is an empty string 


To learn more about regular expressions ("regex"), visit the following sites:

- [regular expression basics](https://www.w3schools.com/python/python_regex.asp).
- [interactive regular expressions tool](https://regex101.com/)

<!-- grade_id: repl, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_repl, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df.columns.tolist() == ['name', 'ID%', 'Height', 'Average Height Parents', 'Gender']
```

```python
df.columns
```

### Replace special characters

- Again, we use regular expressions to deal with special characters (like %, &, $ etc.)

Replace "%" with an empty string

Hint:
    
```python
df.___ = df.columns.str.___(r"___", r"", regex=True)
```

<!-- grade_id: replace_percent, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_replace_percent, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df.columns.tolist() == ['name', 'ID', 'Height', 'Average Height Parents', 'Gender']
```

```python
df.columns
```

### Lowercase and whitespace

We can use two simple methods to convert all columns to lowercase and replace white spaces with underscores ("_"):

- `.str.lower()`


- `.str.replace(' ', '_')`

Hint:

```python
___.___ = ___.___.___.___().___.___(' ', '_')
```

<!-- grade_id: lower, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_lower, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df.columns.tolist() == ['name', 'id', 'height', 'average_height_parents', 'gender']
```

```python
df.columns
```