**Wichtige Hinweise**

- Aktivieren Sie die entsprechende Conda-Umgebung bevor Sie beginnen.

- Geben Sie als `NAME` ihr HdM-Kürzel an.

- Ändern Sie **nicht** den Namen der Datei und löschen Sie keine Zellen.

- Bearbeiten Sie alle Zellen mit dem Hinweis <font color='green'> \# YOUR CODE HERE </font> 

- Die Funktion **NotImplementedError()** soll die Abgabe von leeren Zellen verhindern. Löschen Sie die Funktion, sobald Sie in einer dieser Zellen arbeiten. 

- Stellen Sie sicher, dass alles wie erwartet läuft, bevor Sie die Prüfung abgeben: Starten Sie den Kernel neu und führen Sie alle Zellen aus: wählen Sie "Restart" und dann "Run All" 

```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Selection

## Import pandas

```python
import pandas as pd
```

## Import data

```python
# URL of data
URL = "https://raw.githubusercontent.com/kirenz/datasets/master/height_clean_cols.csv"
```

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Selection

### Getting [[]]

Selecting a single column with `[[]]`:

- Select the column `height` and save it as a new Pandas dataframe `df_height`

Hint:

```python
df_height = df[[___]]
````

<!-- grade_id: select1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_select1, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert  df_height.columns.to_list() == ['height']
```

Selecting via [], which slices the rows (endpoint is not included) and includes all columns:

```python
df[0:1]
```

- Select rows 1, 2 and 3 and save it as `df_sr`

<!-- grade_id: select2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_select2, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert df_sr.iloc[0,0] == 'Peter'
assert df_sr.iloc[2,0] == 'Manuela'
assert len(df_sr) == 3
```

### By label .loc

The `.loc` (location) attribute is the primary access method. The following are valid inputs:

For getting a cross section using a label:

```python
# select row 0
df.loc[[0]]
```

Selecting on a multi-axis by label:

```python
# only select location at row 0 for column "name"
df.loc[0 , 'name']
```

```python
# select row 2 to 4 for column "name" (when using .loc endpoints are included)
df.loc[2:4 , 'name']
```

```python
# select row 2 to 4 for columns "name" and "height" (when using .loc endpoints are included)
df.loc[2:4 , ['name', 'height']]
```

```python
# select all rows for 
df.loc[ : , ["name", "height"]]
```

- Select row 0 and 1 for columns 'name' and 'height' and save it as `df_loc1`

<!-- grade_id: loc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_loc, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert  df_loc1.columns.to_list() == ['name', 'height']
assert  len(df_loc1) == 2
assert  df_loc1.iloc[0,0] == "Stefanie"

```

Reduction in the dimensions of the returned object:

```python
df.loc[0, ["name", "height"]]
```

For getting a scalar value:

```python
df.loc[[0], "height"]
```

### By position .iloc

Pandas provides a suite of methods in order to get purely integer based indexing. 

Here, the `.iloc` attribute is the primary access method. 

*When using `.iloc`, endpoints are not included.*

```python
df.iloc[0]
```

By integer slices:

```python
df.iloc[0:2, 0:2]
```

By lists of integer position locations:

```python
df.iloc[[0, 2], [0, 2]]
```

For slicing rows explicitly:

```python
df.iloc[1:3, :]
```

For slicing columns explicitly:

```python
df.iloc[:, 1:3]
```

For getting a value explicitly:

```python
df.iloc[0, 0]
```

Use `.iloc` to obtain rows 4 to 6 and columns 3 and 4. Save the result as `df_iloc1` 

<!-- grade_id: iloc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_iloc, points: 2, grade: true, solution: false -->
```python
"""Check if your code returns the correct output"""
assert  df_iloc1.columns.to_list() == ['average_height_parents', 'gender']
assert  len(df_iloc1) == 3
assert  df_iloc1.iloc[0,1] == "male"
```