**Important notes:**

- Use your **HdM ID** (e.g. the xy123 in ***xy123***@hdm-stuttgart.de) as **NAME**


- Don't change the name of the file and don't delete any cells.


- Make sure you fill in any place that says  <font color='green'> \# YOUR CODE HERE </font> or "YOUR ANSWER HERE".


- The function `NotImplementedError()` prevents you from hand in tasks with empty cells. Simply delete the function if you start working on a cell with this entry.


- Before you turn this problem in (i.e., after you completed all tasks), make sure everything runs as expected: Restart the kernel and run all cells. If you use:
  - *Visual Studio Code*: select "Restart" and then "Run All" 
  - *Colab*: in the menubar, select `Runtime` and click on `Restart and run all`
  - *Jupyter Notebook*: in the menubar, select `Kernel` and click on `Restart & Run All`


Good luck!

```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Import and store data

## Import pandas

- To load the pandas package and start working with it, import the package. 


- The community agreed alias for pandas is `pd`, so loading pandas as pd is assumed standard practice for all of the pandas documentation:

```python
import pandas as pd
```

## Data creation

- To manually store data in a table, create a DataFrame:

- create the DataFrame and name it `my_df`


- add the variable "gender" with the following observations: male, female and male

Hint:

```python
my_df = pd.DataFrame(
    { 
        'name': [ "Tom", "Lisa", "Peter"],
        'height': [1.68, 1.93, 1.72],
        'weight': [48.4, 89.8, 84.2]
         ______     ________
        }
)
```

<!-- grade_id: my_df, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_my_df, points: 2, grade: true, solution: false -->
```python
"""Check that your code returns the correct output"""
assert my_df.iloc[0, 0] == "Tom"
assert my_df.iloc[0, 3] == "male"
assert my_df.iloc[1, 3] == "female"
```

## Import data

![](https://pandas.pydata.org/docs/_images/02_io_readwrite.svg)

Pandas supports many different file formats or data sources out of the box (csv, excel, sql, json, parquet, …). Each of them import data with the prefix `read_*`

- Next, we import data, available as a CSV file in a GitHub repo 

Important notes: The CSV file has a typical european format with 

- `;` as seperators between cells (instead of the Pandas default `,`)
- `,` as number decimal seperator (instead of the Pandas default `.`) 

We need to take care of this issues.

Hint 

```python
df = pd.___(sep=';', decimal=',')
```

```python
# URL of data
URL = "https://raw.githubusercontent.com/kirenz/datasets/master/height_unclean.csv"
```

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_read_csv, points: 2, grade: true, solution: false -->
```python
"""Check that your code returns the correct output"""
assert df.iloc[0, 0] == 'Stefanie'
assert df.iloc[0, 3] == 161.5
```

Now we do the same import with a different approach. Use the name `df2` 


Hint:

```python
ROOT = "https://raw.githubusercontent.com/kirenz/datasets/master/"
DATA = "___"

df2 = pd.read_csv(___ + ___, sep=";", decimal=',' )

# show head of the data
df2.head()
```

<!-- grade_id: df2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_df2, points: 2, grade: true, solution: false -->
```python
"""Check that your code returns the correct output"""
assert df2.iloc[0, 0] == 'Stefanie'
assert df2.iloc[0, 3] == 161.5
```

## Store data

Pandas supports many different file formats (csv, excel, sql, json, parquet, …). Each of them stores data with the prefix `to_*`


- The following code should save the data `df` as an Excel file in your current directory.


- In the example here, the `sheet_name` is named `people_height` instead of the default Sheet1. 


- By setting `index=False` the row index labels are not saved in the spreadsheet:

```python
df.to_excel("height.xlsx", sheet_name="people_height", index=False)
```

- The equivalent read function `read_excel()` would reload the data to a DataFrame:

```python
# load excel file
df_xlsx = pd.read_excel("height.xlsx", sheet_name="people_height")
```

## Viewing data

### Overview

```python
# show df
df
```

```python
# show first 2 rows
df.head(2)
```

```python
# show last 2 rows
df.tail(2)
```

- The `info()` method prints information about a DataFrame including the index dtype and columns, non-null values and memory usage:

```python
df.info()
```

### Column names

```python
# Show columns
df.columns
```

### Data type

- Show data types ([dtypes](https://pandas.pydata.org/docs/user_guide/basics.html#basics-dtypes)).

```python
df.dtypes
```

- The data types in this DataFrame are integers (int64), floats (float64) and strings (object).

### Index

```python
# Only show index
df.index
```